This assignment is completely voluntary. To check if everything works, we recommend you still upload your finished notebook to moodle and let it be graded.

# Introduction to Jupyter Notebooks

Jupyter Notebooks are a way to combine text and interactive code execution. Every notebook consists of cells like the one you are reading right now. This is a markdown cell for text but there are also code cells like the one below this one. Code cells can be executed by selecting them and then pressing Ctrl+Enter or clicking the run button at the top. Try it with the Cell below! This should output "Hello World!" which will be displayed directly below the executed code cell. 

```python
# This is a code cell. Select it and then execute it by pressing Ctrl+Enter (Strg+Enter). Code cells are python cells by default.
# If you know python try playing around with this cell. This cell will not be graded so change it to your heart's content.
print("Hello World!")
```

## Introduction to nbGrader

Jupyter notebooks are great for providing instructions and incomplete code snippets which you have to complete. While working on assignments you can execute the cells as often as you want, which allows you to play around in a safe environment.

But this is not all: jupyter notebooks allow for automated grading! After completing the given code snippets, you can upload your notebook file as submission to moodle. We will then autograde your submission using nbGrader. This works by executing unit tests on your code, basically checking if your code does what it should do. If a test fails, you will not get points for that task. After autograding you will get a feedback file that shows you which test failed.

To give you a reference, some of the tests will be visible and executable for you. You can use this to check if your code has the correct format. But don't forget that there are also hidden tests, so just because the visible tests work does not mean your code is completely correct.

### Example task

Below is a code cell with incomplete python code. The task is to set variable x to 10 and variable y to 20.
To set a variable in python, you just write its name, an equal sign and then a value. To set a variable z to 30, you would simply write the following:

z = 30

Now try to do the task in the cell below by writing your code just below the "# YOUR SOLUTION HERE" comment.

<!-- grade_id: cell-c7038067c95e57ed, points: 0, grade: false, solution: true -->
```python
# Set x to 10 and y to 20. Don't change code that is already given.
z = 30
### YOUR SOLUTION HERE

print(f"x has the value {x}")
print(f"y has the value {y}")
print(f"z has the value {z}")
```

### Example Test

Below is an example test cell. These cells can not be manipulated but can be executed. This one only tests if x is set to 10 but remember that there are also hidden tests. Execute the test cell to check if your code is at least partially correct. If the cell below fails, you either have an error in your code or forgot to execute your code cell above. Try it now!

<!-- grade_id: cell-95f7b35c0b30bb97, points: 1, grade: true, solution: false -->
```python
assert x==10, f"x has value {x} but should have value 10"
assert z==30, f"z has value {z} but should have value 30. You shouldn't change code that was already given in the snippet."
print("Test passed!")
```

## RDF Extension

We will use an extension to execute and evaluate languages introduced in the Semantic Web lecture. The following cell will be at the top of every assignment and is used to load said extension. You need to run this cell at the start of every session. The cell below the loading cell is a very basic usage of the extension to check whether it works. Run it and it should output a graph.

```python
%reload_ext jupyter-rdfify
```

```python
%%rdf turtle -d graph
@prefix ex: <http://example.org/#>.
ex:YourInstallation ex:status ex:Works.
```

## Submitting your assignment

When you are ready to submit your assignment save the notebook and then upload it to moodle. Try doing the example task and then upload your notebook to moodle. If you did everything right, your submission will be graded and you will get feedback. Also note that all submissions are only graded once after the deadline. You can upload as many tries as you want but only the last submission will be graded.